// variables3.rs
//
// Execute `rustlings hint variables3` or use the `hint` watch subcommand for a
// hint.

// I AM NOT DONE

fn main() {
  let x: i32 = 2600;
  println!("Number {}", x);
}

// let x: i32 indicates that this variable will cover a 32 bits integer
// is a explcit type annotation
// but the variable is not initialized with its value --> let x: i32