// variables2.rs
//
// Execute `rustlings hint variables2` or use the `hint` watch subcommand for a
// hint.

// I AM NOT DONE

fn main() {
  let x = 10;
  if x == 10 {
    println!("x is ten!");
  } else {
    println!("x is not ten!");
  }
}

// 'x' variable can't be initialized as a null --> let x;
// should be initialized like --> let x = number;